Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

```python
# Please add your name and University of Michagan uniqname here...

NAME = 'Ilana Multak'
UMICH_UNIQNAME = 'imultak'
```

---

```python
version = "REPLACE_PACKAGE_VERSION"
```

```python
version = "v2.3.060120"
```

# SIADS 515 Week 3 Homework (HW3)


*The notebook for this assignment is a little different. There are bugs and problems in the existing code that you need to fix in order to get the assert statements to pass. There are no hidden grader cells in this assignment. If you run the first cell under each question title, you’ll either get errors or it will produce the wrong answer. Fix these to get the assert statements to pass. Since there are no hidden assert statements, there’s no real benefit to submitting before you’ve finished your work.*

Think of this assignment in the following way:  it's your first day on the job and you've been given a notebook that was authored by someone who is no longer with your company.  You've been asked to fix it.  There are errors in it, and some of it was not completed by the original author.  You're lucky, though, as there are assertions
sprinkled throughout the notebook to help guide you along the way.


Top-level goal of notebook:
Read a CSV file into a pandas DataFrame and add specific columns to it.
These columns are added by applying functions to specific columns.
The columns to add include:
1. A datetime column that converts "Garmin time" to standard (unix epoch) time.  Note that Garmin doesn't use standard epoch offsets for their timestamps.  Rather than using the number of seconds that have elapsed since midnight on January 1, 1970, they use the number of seconds from midnight on December 31, 1989.

2. A conversion of "semicircles" of latitude and longitude to two different formats: degrees, minutes, seconds 3-tuples and fractional degrees.  For example, a latitude of 504719750 semicircles corresponds to a 3-tuple of degrees, minutes and seconds of (42, 18, 18.43) and 42.305121 degrees.

3. A "normalized speed" column that consists of the values for speed modified to remove outliers by replacing them with upper and lower bounds as well as normalized to z-values (i.e. by subtracting the mean from each value and dividing the result by the standard deviation).

In addition, you will need to complete a function that looks at the difference between sequential rows to determine whether the cyclist is slowing down or not.

Your task for this assignment is to debug this notebook to produce the desired results as shown in the assertions below.

---

## Question 1

```python
import pandas as pd
import numpy as np
ride = pd.read_csv('assets/ride_final2.csv')
ride.head()

def garmin_time_to_datetime(series):
    """Convert Garmin FIT time by adding the number of 
    seconds from January 1, 1970 to December 31, 1989.
    """
    
    return pd.to_datetime(series + 631065600, unit='s', utc=True)
```

<!-- grade_id: cell-2b3a2c26628bb706, points: 0, grade: false, solution: true -->
```python
ride['Timestamp_datetime'] = ride.Timestamp.map(garmin_time_to_datetime)
```

<!-- grade_id: cell-8868f43f56a6d443, points: 10, grade: true, solution: false -->
```python
# This is a read-only grader cell

ride['Timestamp_datetime'] = ride.Timestamp.map(garmin_time_to_datetime)

assert ride.Timestamp_datetime[0] == pd.to_datetime('2018-05-23T14:02:25', utc=True)
```

## Question 2

```python
def semicircles_to_degrees(semicircles):
    '''
    Convert semicircles to degrees
    '''
    max_31_bit_int = 2**31
    return semicircles * (180/max_31_bit_int)


def degrees_to_dms(degrees_fraction):
    ''' Convert degrees to degree, minute, second 3-tuples'''
    degrees = int(degrees_fraction)
    minutes = int((degrees_fraction - degrees) * 60)
    seconds = round((degrees_fraction - degrees - minutes/60) * 3600, 2)
    return (degrees, abs(minutes), abs(seconds))


def dms_to_degrees(d,m,s):
    ''' Convert degrees, minutes, seconds to fractional degrees'''
    return d+m/60+s/3600
```

<!-- grade_id: cell-f5797d720b072f0e, points: 0, grade: false, solution: true -->
```python
dms_to_degrees(*degrees_to_dms(42.2833333333))
```

Output:
```
42.28333333333333
```

```python
dms = degrees_to_dms(42.2833333)
assert dms[1] < 60
```

```python
dms = degrees_to_dms(59.99988)
assert dms[2] < 60
```

```python
ride['Latitude_degrees'] = ride['Latitude'].map(semicircles_to_degrees)
ride['Longitude_degrees'] = ride['Longitude'].map(semicircles_to_degrees)
ride['Latitude_dms'] = ride['Latitude_degrees'].map(degrees_to_dms)
ride['Longitude_dms'] = ride['Longitude_degrees'].map(degrees_to_dms)

last_row = ride.iloc[213]
assert last_row.Latitude_dms == (42, 16, 50.05)
assert last_row.Longitude_dms == (-83,44,21.99)
```

<!-- grade_id: cell-4710d465fe29c816, points: 10, grade: true, solution: false -->
```python
# This is a read-only grader cell

dms = degrees_to_dms(42.2833333)
assert dms[0] >= -180
assert dms[0] <= 180
assert dms[1] >= 0
assert dms[1] < 60
assert dms[2] >= 0
assert dms[2] < 60
assert dms == (42, 16, 59.99988)
assert dms_to_degrees(dms[0], dms[1], dms[2]) == 42.2833333

ride['Latitude_degrees'] = ride['Latitude'].map(semicircles_to_degrees)
ride['Longitude_degrees'] = ride['Longitude'].map(semicircles_to_degrees)
ride['Latitude_dms'] = ride['Latitude_degrees'].map(degrees_to_dms)
ride['Longitude_dms'] = ride['Longitude_degrees'].map(degrees_to_dms)

last_row = ride.iloc[213]
assert round(last_row.Latitude_degrees,6) == 42.280569
assert round(last_row.Longitude_degrees,6) == -83.739442
```

Output:
```
AssertionError                            Traceback (most recent call last)
Input In [104], in <cell line: 9>()
      7 assert dms[1] < 60
      8 assert dms[2] >= 0
----> 9 assert dms[2] < 60
     10 assert dms == (42, 16, 59.99988)
     11 assert dms_to_degrees(dms[0], dms[1], dms[2]) == 42.2833333

AssertionError:
```

## Question 3

```python
def normalize(df, pd_series_name, nsd=2):
    '''
    Take all values that are outside some bound (mean +- 2 sd by default)
    and convert them to the appropriate bound.
    Scale the results to z-scores before returning them
    '''
    df = df.copy()
    pd_series = df[pd_series_name].astype(float)

    # Find upper and lower bound for outliers
    avg = np.mean(pd_series)
    sd  = np.std(pd_series)

    # Calculate the bounds
    lower_bound = avg - nsd*sd
    upper_bound = avg + nsd*sd

    # Collapse in the outliers: replace them with appropriate bound
    df.loc[pd_series < lower_bound , pd_series_name ] = lower_bound
    df.loc[pd_series > upper_bound , pd_series_name ] = upper_bound
    
    return (df[pd_series_name] - avg) / sd
```

<!-- grade_id: cell-01d948a85f3c81e8, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-0c910c3906ac1751, points: 10, grade: true, solution: false -->
```python
# This is a read-only grader cell

ride['Speed_normalized'] = normalize(ride,'Speed')

assert round(ride.iloc[0].Speed_normalized,4) == -1.7737
assert ride.iloc[213].Speed_normalized == -2.0
```

## Question 4

```python
def proportion_slowing(df,series_name):
    ''' Calculate the proportion of rows that represent a slower speed than the previous row'''
    #return 0 # whoops -- ran out of time to do this before I got fired!
    count = 0
    initial_index = df.index[0]
    for i in range(len(df[series_name])):
        if i==0:
            pass
        elif df[series_name][i+initial_index] < df[series_name][(i-1)+initial_index]:
            count +=1
    return count/len(df[series_name])
```

<!-- grade_id: cell-1cc7315a8a08ba64, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-2632f42cc211e6ca, points: 20, grade: true, solution: false -->
```python
# This is a read-only grader cell

assert round(proportion_slowing(ride,'Speed_normalized'),6) == 0.514019
assert round(proportion_slowing(ride[:10],'Speed_normalized'),6) == 0.4
assert round(proportion_slowing(ride[10:],'Speed_normalized'),6) == 0.519608
```

```python

```